# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 4  4  1
 4  3  1
 4  2  3
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 9.0
 8.0
 9.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 4  4  4
 4  3  2
 1  1  3
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 4  4  4
 4  3  2
 1  1  3
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 48  36  20
 36  29  13
 20  13  11
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.619717  0.40685
 0.120604  0.609795
 0.835362  0.207697
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
  7.65860430758066
 11.345581820030544
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.261366   0.261366
 0.861041   0.861041
 0.0691151  0.0691151
```

```julia
?hcat
```

Output:
```
search: [0m[1mh[22m[0m[1mc[22m[0m[1ma[22m[0m[1mt[22m [0m[1mh[22mv[0m[1mc[22m[0m[1ma[22m[0m[1mt[22m [0m[1mh[22mvn[0m[1mc[22m[0m[1ma[22m[0m[1mt[22m Mat[0m[1mh[22m[0m[1mC[22monst[0m[1ma[22mn[0m[1mt[22ms @t[0m[1mh[22mread[0m[1mc[22m[0m[1ma[22mll catc[0m[1mh[22m_ba[0m[1mc[22mktr[0m[1ma[22mce

[36m  hcat(A...)[39m

  Concatenate along dimension 2.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> a = [1; 2; 3; 4; 5][39m
[36m  5-element Vector{Int64}:[39m
[36m   1[39m
[36m   2[39m
[36m   3[39m
[36m   4[39m
[36m   5[39m
[36m  [39m
[36m  julia> b = [6 7; 8 9; 10 11; 12 13; 14 15][39m
[36m  5×2 Matrix{Int64}:[39m
[36m    6   7[39m
[36m    8   9[39m
[36m   10  11[39m
[36m   12  13[39m
[36m   14  15[39m
[36m  [39m
[36m  julia> hcat(a,b)[39m
[36m  5×3 Matrix{Int64}:[39m
[36m   1   6   7[39m
[36m   2   8   9[39m
[36m   3  10  11[39m
[36m   4  12  13[39m
[36m   5  14  15[39m
[36m 
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 6.054570099137534
 6.054570099137534
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.183632  0.25707   0.252801
 0.697986  0.136016  0.800307
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -4.139674407706176
 -0.9797338415813472
 23.563074056363842
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
using LinearAlgebra
```

```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia
dot_v = dot(v, v)
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
outer_v = kron(v,v)
```

Output:
```
9-element Vector{Int64}:
 1
 2
 3
 2
 4
 6
 3
 6
 9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Use [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) to compute the cross product of a vector v with itself and assign it to variable `cross_v`

```julia

```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```